<a href="https://colab.research.google.com/github/ElisaHK/KeplerH/blob/master/Variables_In_Python.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Activity: Variables in Python

## Introduction

In this activity you will practice the following:
- Assigning specific values / types to variables
- Overwriting variables

#### Question 1

Create variable `name` and assign any value you want to it.

<!-- grade_id: cell-89bbfa8c8528e855, points: 0, grade: false, solution: true -->
```python
# Your code here
# name = ...

# your 9code here

```

<!-- grade_id: cell-3e323501e904d29b, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

print('Make sure you created the variable: name')
print(name)


```

#### Question 2

Assign the value of `2.7` to the variable `seconds`.

<!-- grade_id: cell-8d5cb3e863fba294, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here


```

<!-- grade_id: cell-ea1c3340532802f6, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(seconds, float), 'Have you assigned a float value to the variable seconds?'

```

#### Question 3

Assign boolean value `True` to the variable `python_is_fun`.

<!-- grade_id: cell-513f546ddd5fd8a9, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here

```

<!-- grade_id: cell-5ed340946941a67f, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(python_is_fun, bool), 'Have you assigned a boolean value to the python_is_fun variable?'
```

```python
pip install pandas
```

Output:
```
Requirement already satisfied: pandas in /usr/local/lib/python3.10/dist-packages (2.2.2)
Requirement already satisfied: numpy>=1.22.4 in /usr/local/lib/python3.10/dist-packages (from pandas) (1.26.4)
Requirement already satisfied: python-dateutil>=2.8.2 in /usr/local/lib/python3.10/dist-packages (from pandas) (2.8.2)
Requirement already satisfied: pytz>=2020.1 in /usr/local/lib/python3.10/dist-packages (from pandas) (2024.2)
Requirement already satisfied: tzdata>=2022.7 in /usr/local/lib/python3.10/dist-packages (from pandas) (2024.2)
Requirement already satisfied: six>=1.5 in /usr/local/lib/python3.10/dist-packages (from python-dateutil>=2.8.2->pandas) (1.17.0)
```

```python
import pandas as pd
employee = pd.read_csv("https://raw.githubusercontent.com/ddaeducation/DataAnalyst/main/employee_data.csv")

```

```python
employee.tail(50)
```

Output:
```
Name  Years of Employment  Weeks of Vacation  Position
33         James Brown                   18                2.0  Engineer
34  Christopher Harris                    2                6.0    Intern
35          Jane Davis                    9                6.0  Engineer
36        Emily Miller                   20                6.0    Intern
37           Sarah Lee                   17                NaN    Intern
38         James Brown                    1                4.0  Engineer
39       James Jackson                   11                5.0    Intern
40       David Johnson                    8                NaN   Manager
41        Jane Jackson                    5                2.0    Intern
42       Michael Jones                    4                3.0   Analyst
43     Jennifer Harris                   11                6.0    Intern
44         David Moore                    3                5.0   Analyst
45       Sarah Jackson                   16                4.0    Int
```